---
chapter: "none yet"
title: "Title for bee-write-closet5"
author: "Danny Kaplan"
difficulty: "XX"
output: learnr::tutorial
tutorial:
  id: "bee-write-closet5"
  version: 0.1
runtime: shiny_prerendered
date: 2020-07-04 
tags: [R arithmetic, MMAC textbook, 1.1.7.e]
id: bee-write-closet
---

```{r bee-write-closet5-setup, include = FALSE}
library(etude)
library(learnr)
library(gradethis)

knitr::opts_chunk$set(echo = FALSE)
learnr::tutorial_options(exercise.timelimit = 60, 
                 exercise.checker = gradethis::grade_learnr)
```


`r etude::exercise_title()` Evaluate 
$$\left| \pi^2 - \frac{4^3}{2}\right|$$

```{r bee-write-closet5-E1, exercise=TRUE, exercise.cap="Arithmetic in R"}

```

```{r bee-write-closet5-E1-check}
gradethis::grade_result(
  pass_if( ~ .result == abs(pi^2 - (4^3)/2)),
  fail_if( ~ TRUE, message = "Keep trying!")
)
```


