---
title: "An Example Homework"
output: learnr::tutorial
runtime: shiny_prerendered
tutorial:
  id: "submitr-demonstration"
  version: 0.5
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(submitr)
library(QandA)
library(etude)
library(ggformula)
library(mosaic)
library(mosaicCalc)
vfun <- basket::check_valid
knitr::opts_chunk$set(echo = FALSE)
options(tutorial.storage = "none")
learnr::tutorial_options(exercise.timelimit = 60, 
                 exercise.checker = QandA::save_check_inputs)
```

<!-- Start login block -->
```{r __login__, context = "server", echo=FALSE}
# But in this example, we'll replace these to
# make the password and storage local
#storage_actions <- in_local_file("submissions.csv")
storage_actions <- 
        in_google_sheets(
          "1w3fEld2rZlR_6FuzvkA-viOLBA5JdqD3xHl-LuLX3-Y",
          "statprep.annie@gmail.com",
          vfun)
# Local memory

```

`r shiny::textOutput("login_message")` 

`r submitr::login_controls()`

```{r context = "server", echo = FALSE}
submitr::shiny_logic(input, output, session, vfun,
                     storage_actions)
```

```{r context = "server", echo = FALSE}
# observe({
#   if (vfun(input$userID,  "i", input$password)) {
#     output$download <- renderUI(downloadLink("get_submissions",  
#                                              label="Download submissions"))
#   } else {
#     output$download <- renderUI(tagList())
#   }
# }
# )
```

```{r context = "server"}
# valid_id <- reactive({
#   if (!isTruthy(input$userID)) return(FALSE)
#   vfun(input$userID, "u")
# })
# valid_password <- reactive({
#   if (!isTruthy(input$userID)) return(FALSE)
#   if (!isTruthy(input$password)) return(FALSE)
#   vfun(input$userID, "p", input$password)
# })

# set the learnr logging function
# observe({
#   if (valid_id()) {
#     if (valid_password())  {
#       user_identification <- isolate(input$userID)
#     } else {
#       user_identification <- paste("Unauthenticated", isolate(input$userID))
#     }
#     options(tutorial.event_recorder =
#               make_recorder(storage_actions$write,
#                             user_identification))
#             
#   }
# })

# output$login_message <-
#   renderText({ 
#       if (valid_id() && valid_password()) {
#         return("Login SUCCESSFUL.")
#       } 
#       if (valid_id()) {
#         return("Valid ID but password not yet correct")
#       } 
#       if (nchar(input$userID) > 1) {
#         return("Invalid user ID")
#       }
#       return("Please log in.")
#       
#     })
# 
# # Link to download submissions ...
# output$get_submissions <- downloadHandler(
#   filename = function() {
#     paste0("submissions-", Sys.Date(), ".csv")
#   },
#   content = storage_actions$read_submissions
# )
```


<!--End login block-->

## Topic One

```{r planets}
question("Which planet do we live on?",
         answer("Mars", correct=FALSE),
         answer("Earth", correct = TRUE,  message = "Good job Earthling!"),
         answer("Saturn",  correct = FALSE),
         allow_retry = TRUE,
         random_answer_order = TRUE
        )
```

## Topic Two

Some dialog

## Topic Three

```{r cars}
question("Which car  do you drive?",
         answer("Subaru", correct=FALSE),
         answer("Volvo", correct = TRUE,  message = "Good in cold weather."),
         answer("Rav-4",  correct = FALSE),
         allow_retry = TRUE,
         random_answer_order = TRUE
        )
```


## Video

![Just a link](https://youtu.be/zNzZ1PfUDNk)


## Code

```{r ash-send-shirt-C1, exercise=TRUE, exercise.cap="Sandbox"}

```

```{r ash-send-shirt-C1-check}
gradethis::grade_result(
  pass_if( ~ .result == 4, message = "Four is a charm!"),
  fail_if( ~ .result == 3, message = "Not three!"),
  pass_if( ~ .result == 2, message = "Two is also good."),
  fail_if( ~ TRUE, message = "You're out of line.")
)
```
