---
title: "An Example Homework"
output: learnr::tutorial
runtime: shiny_prerendered
tutorial:
  id: "submitr-demonstration"
  version: 0.5
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(submitr)

library(etude)
library(ggformula)
library(mosaic)
library(mosaicCalc)
knitr::opts_chunk$set(echo = FALSE)

```

<!-- Start login block 
# Another possibility
storage_actions <- in_local_file("submissions.csv")
-->

```{r} 
submitr::login_controls()
```


```{r context = "server", echo = FALSE}
options(tutorial.storage = "none")
learnr::tutorial_options(exercise.timelimit = 60, 
                 exercise.checker = gradethis::grade_learnr)
vfun <- basket::check_valid
storage_actions <- 
        record_gs4(
          "1w3fEld2rZlR_6FuzvkA-viOLBA5JdqD3xHl-LuLX3-Y",
          "statprep.annie@gmail.com",
          vfun)
submitr::shiny_logic(input, output, session, vfun,
                     storage_actions)
```

<!--End login block-->

`r include_etude("../From_math_135/Functions/ash-send-shirt.Rmd")`

`r include_etude("../From_math_135/Functions/beech-sleep-drawer.Rmd")`

`r include_etude("../From_math_135/Functions/cat-lend-futon.Rmd")`




## Topic One

```{r planets}
question("Which planet do we live on?",
         answer("Mars", correct=FALSE),
         answer("Earth", correct = TRUE,  message = "Good job Earthling!"),
         answer("Saturn",  correct = FALSE),
         allow_retry = TRUE,
         random_answer_order = TRUE
        )
```

## Topic Two

Some dialog

## Topic Three

```{r cars}
question("Which car  do you drive?",
         answer("Subaru", correct=FALSE),
         answer("Volvo", correct = TRUE,  message = "Good in cold weather."),
         answer("Rav-4",  correct = FALSE),
         allow_retry = TRUE,
         random_answer_order = TRUE
        )
```


## Video

![Just a link](https://youtu.be/zNzZ1PfUDNk)


## Code

```{r just_some_code-C1, exercise=TRUE, exercise.cap="Sandbox"}

```

```{r just_some_code-C1-check}
gradethis::grade_result(
  pass_if( ~ .result == 4, message = "Four is a charm!"),
  fail_if( ~ .result == 3, message = "Not three!"),
  pass_if( ~ .result == 2, message = "Two is also good."),
  fail_if( ~ TRUE, message = "You're out of line.")
)
```
